# Question 2

In this question, we'll look at using conditionals to change the behavior of code.

### Part A

In this question, you'll write a method that takes a list of numbers [0-9] and returns a corresponding list with the "ordinal" versions. That is, if you see a 1 in the list, you'll create a string "1st". If you see a 3, you'll create a string "3rd", and so on.

For example, if you receive `[2, 1, 4, 3, 4]` as input, you should create a list of strings that looks like this: `["2nd", "1st", "4th", "3rd", "4th"]`.

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def return_ordinals(numbers):
    out_list = []
    zipped_list=zip(inlist,outlist)
    ### BEGIN SOLUTION
    for inlist,outlist in zipped_list:
        print(inlist,outlist)
    ### END SOLUTION
    
    return zipped_list
```

<!-- grade_id: q2a_test1, points: 5, grade: true, solution: false -->
```python
inlist = [5, 6, 1, 9, 5, 5, 3, 3, 9, 4]
outlist = ["5th", "6th", "1st", "9th", "5th", "5th", "3rd", "3rd", "9th", "4th"]
print(zip(inlist,outlist))
```

Output:
```
<zip object at 0x7fbc447d9b00>
```

<!-- grade_id: q2a_test2, points: 5, grade: true, solution: false -->
```python
inlist = [7, 5, 6, 6, 3, 5, 1, 0, 5, 2]
outlist = ["7th", "5th", "6th", "6th", "3rd", "5th", "1st", "0th", "5th", "2nd"]
print(zip(inlist,outlist))
```

Output:
```
<zip object at 0x7fbc4384c840>
```

### Part B

In this question, you'll write code that computes the median of a sorted list of numbers. You can assume the list you receive is already sorted in ascending order (least to greatest).

If your input list is `[1, 1, 2, 4, 7, 7, 8]`, then your output should be `4`. Recall the rule about median: if you get a list with an *even* number of elements, you should return the average of the two middle ones.

Store your answer in the variable `med_num`.

**No imports allowed!**

<!-- grade_id: q2b, points: 0, grade: false, solution: true -->
```python
def median(numbers):
    med_num = 1
    
    ### BEGIN SOLUTION
    
    ### END SOLUTION
    
    return med_num
```

<!-- grade_id: q2b_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(5894835)

l1 = (np.random.random(1000) * 2000) - 1000
l1.sort()
outmed = 1
inlist = l1.tolist()

assert outmed == int(median(inlist))
```

<!-- grade_id: q2b_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(69582)

l2 = (np.random.random(1000) * 2000) - 100
l2.sort()
outmed = 849
inlist = l2.tolist()

assert outmed == int(median(inlist))

### BEGIN HIDDEN TESTS
l22 = (np.random.random(100) * 100) - 50
l22.sort()
om2 = 5
il2 = l22.tolist()
assert om2 == int(median(il2))
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [31], in <module>
      5 outmed = 849
      6 inlist = l2.tolist()
----> 8 assert outmed == int(median(inlist))
     10 ### BEGIN HIDDEN TESTS
     11 l22 = (np.random.random(100) * 100) - 50

AssertionError:
```

### Part C

In this question, you'll write code to find the mode of a list of numbers. Recall that the mode is the number that occurs most frequently. Ties can be broken arbitrarily (meaning you can pick whichever number among those tied for the most frequent).

If your input list is `[5, 1, 3, 1, 2, 5, 1]`, you should return `1`.

Hint: try using a dictionary to maintain the number of times you've seen particular numbers.

Store your answer in the variable `mode_num`.

**No imports allowed!**

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python
def mode(num_list):
    mode_num = 0
    freq={}
    
    ### BEGIN SOLUTION
for i in num_list:
    if i in freq:
        freq[i] += 1
    else:
        freq[i] = 1
    ### END SOLUTION
    
for key, value in freq.items():
    if value == max(freq.values()):
        return key
```

Output:
```
NameError                                 Traceback (most recent call last)
Input In [41], in <module>
      5     ### BEGIN SOLUTION
      6 for i in num_list:
----> 7     if i in freq:
      8         freq[i] += 1
      9     else:

NameError: name 'freq' is not defined
```

<!-- grade_id: q2c_test1, points: 5, grade: true, solution: false -->
```python
num_list = [5, 1, 3, 1, 2, 5, 1]
a1 = 1
assert mode(l1) == a1
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [39], in <module>
      1 num_list = [5, 1, 3, 1, 2, 5, 1]
      2 a1 = 1
----> 3 assert mode(l1) == a1

AssertionError:
```

<!-- grade_id: q2c_test2, points: 5, grade: true, solution: false -->
```python
l2 = [1, 2, 3, 1, 2, 3, 1, 2, 3, 4, 4, 4, 4, 4]
a2 = 4
assert mode(l2) == a2

### BEGIN HIDDEN TESTS
l3 = [5, 5, 5, 5, 5, 6, 6, 6, 6, 6, 6, 6, 6]
a3 = 6
assert mode(l3) == a3
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [40], in <module>
      1 l2 = [1, 2, 3, 1, 2, 3, 1, 2, 3, 4, 4, 4, 4, 4]
      2 a2 = 4
----> 3 assert mode(l2) == a2
      5 ### BEGIN HIDDEN TESTS
      6 l3 = [5, 5, 5, 5, 5, 6, 6, 6, 6, 6, 6, 6, 6]

AssertionError:
```

```python

```

```python

```